Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tiago Flora"
COLLABORATORS = ["Ara Mkhoyan", "Chloe Gabrielle"]
```

---

# CS110 Pre-class Work 1.2



## Question 1
The following pseudocode has been extracted from Cormen et al. You should analyze the lines of the pseudocode trying to understand what every line in the pseudocode does and how many times is that line executed. 
![Insertion sort pseudo-code](insertionsort_pseudo.png)

Answer the following questions: What does line 7 do? Why is it necessary?

<!-- grade_id: cell-3bac59ff24b11a27, points: 0, grade: true, solution: true -->
Line 7 guarantees two things that are fundamental for insertion sort. It decreases the value of the index i in every iteration of the while loop. That means it is a vital line to meet the termination condition of the while loop, which is executed when i is positive, otherwise it would be executed until a runtime error shows up or your computer crashes. The other result of this is to 'move' the elements A[i] along the array, to leave space for A[j] to be inserted in the correct slot.

## Question 2
The following Python code is an attempt to implement the Insertion-Sort pseudocode above. However, this code has a bug. 


```python
def insertionSort(A):
    for j in range(len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
            A[i+1] = key
    return A
```

Copy and paste the code in the cell below. Identify and correct the bug. 

<!-- grade_id: cell-1b7488e89cb501c5, points: 0, grade: false, solution: true -->
```python
def insertionSort(A):
    for j in range(len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
        A[i+1] = key
    return A
```

<!-- grade_id: cell-3432ce8ae2493fa6, points: 1, grade: true, solution: false -->
```python
assert(insertionSort([0]) == [0])
assert(insertionSort([-1,1]) == [-1,1])
assert(insertionSort([1,-1]) == [-1,1])
assert(insertionSort([1,6,3,6]) == [1,3,6,6])
```

How do you know this code works as intended now?

<!-- grade_id: cell-442998d5fdd8b561, points: 0, grade: true, solution: true -->
If we run the cell with all the 'assert' statements, we see that it does nothing and there is no output. Before, as the assert conditions evaluated to false, we had four AssertionError instances. If we check the comparisons between the function's outpupt and the sorted arrays we input into it, we find they're all true, as expected. <br>
Another way to know the code works as intended now is that removing the indentation from 'A[i+1]=key' guarantees the maintenance of the 'while' loop invariant, which in this case is that every element of A[1..i] is smaller than A[j].

```python
print(insertionSort([0]) == [0])
print(insertionSort([-1,1]) == [-1,1])
print(insertionSort([1,-1]) == [-1,1])
print(insertionSort([1,6,3,6,4]) == [1,3,4,6,6])
```

Output:
```
True
True
True
True
```

## Question 3 
Implement the bubble sort algorithm in the cell below. First, look up the corresponding pseudocode in Cormen et al. (page 40), and write a Python implementation for the bubble sort.  Like in the `insertionSort`,  your new Python function, `bubbleSort`, takes a list of elements as an argument and returns the sorted list.

<!-- grade_id: cell-b709733c96d8b615, points: 0, grade: false, solution: true -->
```python
def bubbleSort(A):
    for i in range(len(A)):
        # j must start at len(A)-1, and  move backwards
        for j in range(len(A)-1, 0, -1):
            # If an element is larger than the preceding number, they swap orders
            if A[j] < A[j-1]:
                A[j], A[j-1] = A[j-1], A[j]
            j -= 1
    return A
```

<!-- grade_id: cell-ffc9b251ea0047ba, points: 2, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

Is your program bug-free? Think of three different [“test cases”](https://www.hiredintech.com/classrooms/algorithm-design/lesson/83) to check the correctness of your code. If a code passes the tests you have just outlined, is this an unambiguous proof that the code is correct? Give your answer in the cell below. Feel free to use additional cells (Markdown or Code) as you wish.

<!-- grade_id: cell-5ecb51404140cb2f, points: 0, grade: true, solution: true -->
Given the definition of the sorting problem, the program seems to be bug-free. Testing a few edge cases, we can further strenghten the point. We tested the program with negative numbers, repeated numbers, a single-element list, a sorted list, and even strings. All the numerical tests return correctly sorted lists, and the strings are returned in alphabetical order (with upper-case letters having priority over lower-case letters). <br>
It's not unambiguous proof the code is correct. We'd need 

```python
print(bubbleSort([3,2,-3,1,5,4,-4,-2,10,8]), 
      bubbleSort([0,1,0,2,0,2,0]),
      bubbleSort([1]),
      bubbleSort([4,5,9,100]),
      bubbleSort(["This", "is", "absolutely", "Amazing"]))
```

Output:
```
[-4, -3, -2, 1, 2, 3, 4, 5, 8, 10] [0, 0, 0, 0, 1, 2, 2] [1] [4, 5, 9, 100] ['Amazing', 'This', 'absolutely', 'is']
```

## Question 4
The following pseudocode corresponds to a third sorting algorithm to be covered in Lesson 1.2, the selection sort.  Like in the previous step, write a Python implementation of the algorithm and make sure that your program is bug-free. `selectionSort` should take a list of elements as an argument and return the sorted list.

![Selection sort pseudo-code](selectionsort_pseudo.png)

<!-- grade_id: cell-e04f94e5d418923c, points: 0, grade: false, solution: true -->
```python
# Following the pseudo-code provided in Cormen, we get:
def selectionSort(A):
    # Iterate over the whole array and assume an element is the smallest for comparison
    for i in range(len(A)):
        minidx = i
        for j in range(len(A)):
            # For every next element, compare it to the current one
            if A[j] > A[minidx]:
                # If any element is smaller, the minimum element is now assumed to be it
                minidx = j
            # Swap the positions of the original element being compared (A[i]) and the new minimum
            A[i], A[minidx] = A[minidx], A[i]
    return A
```

```python
print(selectionSort([-2,1,5,3,-4,6,3,10]))
```

Output:
```
[-4, -2, 1, 3, 3, 5, 6, 10]
```

<!-- grade_id: cell-ffe912ac7922de70, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

**Please make sure you try your best to code insertion sort, selection sort and bubble sort, and bring these programs to class. You will need them for a class activity.**