Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tiago Flora"
COLLABORATORS = ["Ara Mkhoyan", "Chloe Gabrielle", "Cormen et al and the internet"]
```

---

# CS110 Pre-class Work 2.2

## Question 1 (Exercise 3.1-3 of Cormen, et al. )
Explain why the statement, "The running time of algorithm A is at least $O(n^2)$," is meaningless.


<!-- grade_id: cell-437d61c1420d4f99, points: 0, grade: true, solution: true -->
The key of the issue with the statement lies on the words "at least." <br>
Big-Oh (O-) notation denotes the set of functions that express the maximum running time of an algorithm of input size n. Thus, when we say the running time $T(n)\geq g(n)$, and $g(n)$ is a function in $O(n^2)$, we are comparing completely arbitrary terms. It could be that the function $g(n)$ used as reference is a constant, which would be $O(n^2)$, but would be obvious for large enough input sizes. In the clearest example, if we just said $T(n)$ is positive ($g(n)=0$), it states the obvious, and so we see that the statement is redundant.

## Question 2 (Exercise 3.1-4 of Cormen, et al. )

Is $2^{n+1}=O(2^n)$? Is $2^{2n}=O(2^n)$?

<!-- grade_id: cell-6a0c4ee19dadfddf, points: 0, grade: true, solution: true -->
### 1. $2^{n+1}$ 

Yes. To check whether $T(n)=2^{n+1}=O(2^n)$, we need to prove $0 \leq 2^{n+1} \leq c2^{n}$ for any $n\geq n_0$. Since $2^n$ will be positive for any real value of n, we worry about the right side of the inequality. First, in the case $n=n_0$, we have:
$$2^{n_0+1} = c2^{n_0} \implies 2\cdot 2^{n_0} = c2^{n_0}$$
Since $2^{n_0}\neq0$, we can divide both sides by it and we arrive at $c=2$. $n_0$, thus, is the smallest of the possible values it can assume. Given that $n_0 > 0$ is the only bound on its value, and $n_0$ is an integer, we have $n_0=1$. Thus, we conclude $T(n)=2^{n+1}=O(2^n)$.

### 2. $2^{2n}$ 

No. Again, we need to prove $0 \leq 2^{2n} \leq c2^{n}$. We have:
$$2^{2n} = 2^{n}\cdot2^{n} \leq c2^{n}$$ 
$$ \therefore 2^{n}\leq c $$
It is impossible to get a constant that is asymptotically larger than $2^n$. Thus, $2^{2n} \neq O(2^n) $


## Question 3.
Write a function in Python that solves the maximum-subarray problem using a brute-force approach. Your Python function must:
* Take as Input an array/list  of numbers
* Produce the following Output: 
    * the start and end indices of the subarray containing the maximum sum.
    * value of the maximum subarray (float)


<!-- grade_id: cell-527e6532d6992fd0, points: 0, grade: false, solution: true -->
```python
import numpy as np
import math

def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    maximum = -float(math.inf)
    right_max = 0
    left_max = 0
    for i in range(len(A)):
        # Every iteration of the loop will restart the sum of the subarray
        subsum = 0
        for j in range(i, len(A)):
            subsum += A[j]
            """
            Instead of changing the maximum subarray inside the if statement,
            we can define it as max(subsum, maximum).
            However, if we do this, we will get an AssertionError for the first test,
            for the algorithm will choose to update the starting index of
            the maximum subarray of the first assert test, which can be either 1 or 3.
            """
#            maximum = max(subsum, maximum)
            if maximum < subsum:
                maximum = subsum
                right_max = j
                left_max = i
    return((left_max, right_max, maximum))
```

<!-- grade_id: cell-a28a56466c9537ff, points: 1, grade: true, solution: false -->
```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

```python
print(bruteforce_max_subarray([-2,1,-1,2,-5]), # The algorithm chooses A[1] as the starting point
bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]))
```

Output:
```
(1, 3, 2) (2, 6, 7)
```

## Question 4. 
Test your Python maximum-subarray function using the following input list (from Figure 4.3 of Cormen et al.):  
`A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7] `

If your Python implementation is correct, your code must return: 
* 43 - which is the answer to the maximum subarray problem, and 
* <7, 10> -the start and the end indices of the max subarray. 



<!-- grade_id: cell-e4a632ce0edc1697, points: 0, grade: true, solution: true -->
```python
A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]
bruteforce_max_subarray(A)
```

Output:
```
(7, 10, 43)
```

## Question 5. Asymptotic notation. 
Complete the following table using the asymptotic notation that best describes the problem. For example, if both $O(n^3)$ and $O(n)$ are possible for an algorithm, the answer is $O(n)$ because the function $f(n) = O(n)$ provides a tighter and more accurate fit; if both $O(n)$ and $\Theta(n)$ are possible, the correct answer is $\Theta(n)$ because $\Theta(n)$ provides both information about the upper and lower bound, thus it is more accurate than $O(n)$.

You should copy the following table and paste and edit it in the cell below. 

Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort |  |
Selection sort |  |
Bubble sort |  | 
Finding maximum subarray |  |

<!-- grade_id: cell-247c51df276b622e, points: 0, grade: true, solution: true -->
Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort |  | $\Theta(n^2)$
Selection sort |  | $\Theta(n^2)$
Bubble sort |  | $\Theta(n^2)$
Finding maximum subarray | $O(n^2)$ | 

## [Optional] Question 6. 
How can you change this code to make it find the minimum-subarray?

<!-- grade_id: cell-f4553f71858d1bc4, points: 0, grade: false, solution: true -->
```python
def bruteforce_min_subarray(A):
    """
    Implements brute-force minimum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the min subarray
    - the end index of the min subarray
    - the value of the min subarray
    """
    minimum = float(math.inf)
    right_min = 0
    left_min = 0
    for i in range(len(A)):
        # Every iteration of the loop will restart the sum of the subarray
        subsum = 0
        for j in range(i, len(A)):
            subsum += A[j]
#            minimum = min(subsum, minimum)
            if minimum > subsum:
                minimum = subsum
                right_min = j
                left_min = i
    return((left_min, right_min, minimum))
```

<!-- grade_id: cell-4263c7494a5f09d3, points: 0, grade: true, solution: false -->
```python
assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
assert(bruteforce_min_subarray([1]*10)[2] == 1)
```

```python
bruteforce_min_subarray(A)
```

Output:
```
(1, 6, -50)
```