Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tiago Flora"
COLLABORATORS = ["Ara Mkhoyan","Chloe Gabrielle"]
```

---

# CS110 Pre-class Work 3.2

## Question 1.
Given the array `H=[39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]`, perform the following operations:
1. Draw the corresponding binary tree of H. Is the binary tree a valid max heap? Explain your answer.
2. Using as a model the drawing examples illustrated in Figure 6.2 of Cormen et al.,  draw a step-by-step transformation of the array above into a valid max heap. 
3. Now that you have obtained a valid max heap, write out the corresponding array that stores the valid max-heap.

Use as many cells as you wish for this question.

<!-- grade_id: cell-b1f821b034b44619, points: 0, grade: true, solution: true -->
## 1.
The binary tree of H is not a valid max heap. It doesn't hold the max-heap property that every parent node should be larger than its children.

![img](https://i.ibb.co/SsCbTRd/CS110-3-2-1.jpg)

## 2.

![img](https://i.ibb.co/9YYRMkd/CS110-3-2-2.jpg)
![img](https://i.ibb.co/pRnG7xt/CS110-3-2-3.jpg)

## 3.
The max-heap version of the array H can be seen in step $k$ of the steps above.

```python
H = [92,85,85,76,49,30,49,16,39,15,21,7,29,31,28]
```

## Question 2. 
Consider the following questions on the $MAX-HEAPIFY$ operation.
### Question 2a.

In the pseudocode of $MAX-HEAPIFY$ (Cormen et al., p.154, or you can view it [here](https://drive.google.com/open?id=1e_3jsX4-qQCfZXKMok_T6LPFh9FwtmT5)), what does A.heap-size mean and what is the idea behind the local variable largest? 


<!-- grade_id: cell-06106b81a909b9e6, points: 0, grade: true, solution: true -->
To any array that represents a heap, its heap-size refers to the amount of elements in the heap that are stored within the array. That is why it is used to check if elements to be operated on in the max-heapify algorithm are within the heap. <br>
The variable "largest" is used to store the maximum value in a set of three nodes: one parent, and its two children. It is used to store the maximum value in the trio that should be moved to the lowest index value.

### Question 2b.
The functions $LEFT(i)$ and $RIGHT(i)$, lines 1 and 2 in the $MAX-HEAPIFY$ pseudocode, return the array index of the left and right child, respectively, of a node in a binary tree. From reading Section 6.1, you know that the input to both functions is an integer number, $i$, which corresponds to the array index of the parent node in the array. Review Section 6.1 for more information. Write a Python implementation of the functions $LEFT(i)$ and $RIGHT(i)$ by filling in the cells below.

<!-- grade_id: cell-2efd4017321c8d48, points: 0, grade: false, solution: true -->
```python
def left(i):
    # Returns the index of the left child node
    return(2*i+1) # Because Python indexes from 0, we add 1 to 2i to get the left child's index
```

<!-- grade_id: cell-507ddbab62e6a32c, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

<!-- grade_id: cell-b7e5ceeacaeaf1b8, points: 0, grade: false, solution: true -->
```python
def right(i):
    # Returns the index of the right child node
    return(2*i+2) # To get the right index, we add 2 to 2i
```

<!-- grade_id: cell-c24867515653788a, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 2c.
Write a Python implementation of the MAX-HEAPIFY operation using the pseudocode above, and your newly written functions, `left` and `right`.

<!-- grade_id: cell-52289253a243341b, points: 0, grade: false, solution: true -->
```python
def heapify(heap, i):
    """
    Inputs:
    - heap: a list of floats. Assume that the heap size is the length of the heap.
    
    No output is needed. This function should modify (if necessary) heap in-place.
    """
    # We subtract 1 from the heap_size to avoid index errors
    heap_size = len(heap)-1 # Alternatively, we could just use a strictly smaller (< instead of <=) comparison operator with r and l
    l = left(i) # Define the left child node's index
    r = right(i) # Define the right child node's index
    if l <= heap_size and heap[l] > heap[i]: # Check if left child is in the heap and if it is larger than its parent
        largest = l
    else: # Keep parent node as largest
        largest = i
    if r <= heap_size and heap[r] > heap[largest]: # Check if right child is in the heap and if it's larger than its parent
        largest = r
    if largest != i: # In case the parent node does not contain the largest element, swap elements with the child node with the largest element
        heap[i], heap[largest] = heap[largest], heap[i]
        heapify(heap, largest) # Recursively call heapify until the parent is the largest node
```

<!-- grade_id: cell-7f63745ff34b881b, points: 1, grade: true, solution: false -->
```python
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
heapify(A,0)
assert(A == [85, 49, 85, 16, 39, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28])

```

## Question 3. 
Next, write a Python implementation of the BUILD_MAX_HEAP operation using the pseudocode provided in Section 6.3 of Cormen et. al. Test your Python implementation using the array in problem 1, and make sure your Python codes produce a valid max heap.

<!-- grade_id: cell-9237da6268a5e3eb, points: 0, grade: false, solution: true -->
```python
import math

def build_max_heap(A):
    """
    Input:
    - A: a list of floats.
    
    No output is needed. The function should turn A into a valid max heap, in-place.
    """
    for i in range(math.floor(len(A)/2), -1, -1): # We start from the highest-indexed non-leaf node and heapify from there until the first element
        heapify(A, i) # "Heapifies" every branch of the heap
```

<!-- grade_id: cell-5dc147b837da3f7e, points: 1, grade: true, solution: false -->
```python
A = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
build_max_heap(A)
assert(A == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1])
```

## Question 4. 

Lastly, write Python implementations of the $MIN-HEAPIFY$ and $BUILD-MIN-HEAP$ operations for a min heap data structure. You can use your $MAX-HEAPIFY$ and $BUILD-MAX-HEAP$ Python function as models, just remember that the latter two functions support operations for the max heap data structure. Test your Python implementation of $BUILD-MIN-HEAP$ using the array in problem 1, and make sure your Python codes produce a valid min heap. 

<!-- grade_id: cell-d86943f127876fc8, points: 0, grade: false, solution: true -->
```python
def min_heapify(heap, i):
    """
    Inputs:
    - heap: a list of floats. Assume that the heap size is the length of the heap.
    
    No output is needed. This function should modify (if necessary) heap in-place.
    """
    # The same rationale for max_heapify applies here for the indices
    heap_size = len(heap)-1
    l = left(i)
    r = right(i)
    if l <= heap_size and A[l] < A[i]: # We now check if a child node is smaller than its parent
        smallest = l
    else:
        smallest = i
    if r <= heap_size and heap[r] < heap[smallest]:
        smallest = r
    if smallest != i: # If the smallest is not the parent, swap elements and call the function recursively
        heap[i], heap[smallest] = heap[smallest], heap[i]
        min_heapify(heap, smallest)
```

<!-- grade_id: cell-6efa45e3956ad0b8, points: 0, grade: false, solution: true -->
```python
def build_min_heap(A):
    """
    Input:
    - A: a list of floats.
    
    No output is needed. The function should turn A into a valid min heap, in-place.
    """
    for i in range(math.floor(len(A)/2), -1, -1): # We iterate from the last parent node up to the top one
        min_heapify(A, i)
```

```python
A = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
build_min_heap(A)
A
```

Output:
```
[1, 2, 3, 4, 7, 9, 10, 14, 8, 16]
```

<!-- grade_id: cell-7e64941402f819c2, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```