Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tiago Flora"
COLLABORATORS = ["Chloe Go", "Hanaka Saito"]
```

---

# CS110 Pre-class Work 5.1

## Question 1.

### Question 1a.

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
import random

def worst_sort(A):
    """
    Sort A in ascending order by randomly shuffling its 
    elements until they are in order.
    
    Input:
    - A: list of numerical values
    
    Output:
    - A: Sorted list
    """
    def A_is_sorted(A): # Define a local function that will compare elements 
        # of a list and return a boolean indicating if it is sorted
        
        for j in range(len(A)-1):
            
            if A[j] > A[j+1]: # Compare A[j] to its next neighbor
                return False # Return False if any elements aren't sorted
            
        return True # Not passing any if statement over the iterations, return True as the list is sorted
    
    while A_is_sorted(A) == False: # Keep shuffling A until it becomes sorted by chance (!)
        random.shuffle(A)
    return A
```

```python
A = [-10,-4,-11,5,1,4,0,5,2,-3,-6]

worst_sort(A)
```

Output:
```
[-11, -10, -6, -4, -3, 0, 1, 2, 4, 5, 5]
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b.
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->
In the best case, the list passed to the algorithm is already sorted, and thus the only procedure the algorithm performs is comparing all the elements of the list to its next neighbor to confirm the list is in ascending order. Thus, the best case complexity of the algorithm is $O(n)$, as in that case it depends linearly on how many elements there are to compare.

### Question 1c.

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
The average-case complexity of the algorithm depends on the probability of getting the sorted list in a random permutation of it. <br>
We know that there are $n!$ permutations of a list with $n$ elements, and so if the probability disitribution of each possible permutation is uniform, we can assumed that after any given random shuffle the prbability of it being sorted is $\frac{1}{n!}$. <br>
Thus, we can expect the algorithm to run $n!$ shuffles, on average, to find a correct, sorted list from shuffling. However, for every run, there still are the comparisons to be made between neighboring elements, linearly dependent on $n$. Thus, the running time $T(n) = cn\cdot n!$. Thus, we could say the average case complexity is $O(n\cdot n!)$

### Question 1d.

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
It obviously depends largely on hardware, but even so list sizes couldn't be larger than a few tens of elements. Because time grows in factorial order, running times already become larger than a minute on average laptops for list sizes $n=11$ and higher, for instance.

### [Optional] Question 1e.

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2.

Approximate median finder. Given a list and δ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the (50-tilde{δ}/2)% and (50+\tilde{δ}/2)% percentiles, where \tilde{δ}=100*δ. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a.
Complete the following function


<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
def check_approx_med(A, num, delta):
    """
    Given a list and a number in a list, determine whether it is within (50-delta/2)% and 
    (50+delta/2)% percentiles.
    X% percentile of a list is defined to be the smallest number in the list that is as large as at least
    X% numbers in the list. 
    
    Input:
    - A: list of numerical values
    - num: the number of interest
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - isApproxMed: a boolean value indicating whether num is within the bound
    """
    
    if len(A) == 1: # In the case when A contains a single element, we cannot divide it into percentiles
        isApproxMed = False
        return isApproxMed
    
    def bubbleSort(A):
        for i in range(len(A)):
            # j must start at len(A)-1, and  move backwards
            for j in range(len(A)-1, 0, -1):
                # If an element is larger than the preceding number, they swap orders
                if A[j] < A[j-1]:
                    A[j], A[j-1] = A[j-1], A[j]
                j -= 1
        return A
    
    # Sort A with some sorting method (such as sorted() or bubble sort)
    bubbleSort(A)
    
    # Find the index of the highest element in the lower percentile of the range
    perclow_i = int((50-(delta*100)/2)*len(A)/100)
    
    # Find the index of the lowest element in the higher percentile of the range
    perchigh_i = int((50+(delta*100)/2)*len(A)/100)
    
    # If those elements are the same to the number of interest, 
    if A[perchigh_i] >= num and A[perclow_i] <= num:
        isApproxMed = True
    else:
        isApproxMed = False
    return isApproxMed
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0,1], 0, .25) == True)
assert(check_approx_med([0], 0, .25) == False)
```

### Question 2b.
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
def approx_med_finder(A, delta):
    """
    Given a list, find a number in the list that is between 50-delta/2% and 50+delta/2% percentiles
    
    Input:
    - A: list of numerical values
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - num: the approximated median, if it is found within 100//delta trials (this is one possible 
    maximum number of retries. While you are encouraged to play around with another limits, the 
    submitted work must use this number.)
    - None: finding failed, if nothing found within 100//delta trials

    Note: the 100//delta is chosen here as it is the expected number of trials for the first successful finding to occur (using geometric distribution)
    """
    if len(A) == 1: # Again we check if A contains more than 1 element
        return False
    
    times = int(100//delta) # We define the number of trials as stated in the function's description
    
    for i in range(times): # Choose a random element of A until it is inside the interval of interest
        
        # Choose a random element of A
        num = random.choice(A) 
        
        if check_approx_med(A, num, delta) == True:
            return num
    return None
```

```python
A = [-10,-4,-11,5,1,4,0,5,2,-3,-6,2,5]
delta = 0.2

approx_med_finder(A, delta)
```

Output:
```
1
```

### Question 2c.

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with δ or N (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
The probability of failure depends on $\delta$. For each trial, we know that in $100-\tilde{\delta}\%$ of possible values, the function fails. That probability equates, then, $P(failure) = \frac{100-\tilde{\delta}}{100}$. After all allowed random trials, we will have compounded that probability to become $P(failure | m trials) = \left(\frac{100-\tilde{\delta}}{100}\right)^m$. Both of these probabilities, then, scale with $\delta$, as the absolute size of the list doesn't matter, but the proportion of it inside the interval delta defines does.

### Question 2d.
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->
The running time of check_approx_med depends only on the sorting algorithm used inside of it. In our case, we used bubble sort, an arbitrary decision based off previous needs to use some sorting algorithm, which has an expected runtime on the order of $O(n^2)$. All the other operations in the function are comparisons of time $O(1)$. <br>
In the case of approx_med_finder, we apply check_approx_med a maximum of $100/\delta$ times. Thus, we'd expect the running time to be inversely proportional to $\delta$. However, given that probability of success increases with $\delta$, exiting the loop, both factors even each other out. Then, approx_med_finder has the form $O(n^2)$.